# COMP SCI 1015 IAP - W01 - Workshop submission

## Submission Exercise
This is your first workshop submission. Please write the code that print out `The computer says "Hello World!"`.

<!-- BEGIN QUESTION -->



```python
# INSERT YOUR CODE BELOW
print("HELLO,WORLD");
```

Output:
```
HELLO,WORLD
```

<!-- END QUESTION -->

## <font color='red'>Remember to submit your work to Gradescope and sign out with yout tutor before leaving the session.</font>

```python

```

```python

```

Tests for Task1 (points: 100)
```
>>> out = get_print_output(); # retrieves the print output.
>>> reset_print_output(); # Resets the print output.
>>> 'Hello World' in out
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
